
# Exercise Sheet No. 1

---

> Machine Learning for Natural Sciences, Summer term 2026, Prof. Pascal Friederich, pascal.friederich@kit.edu
> 
> Tutor: jonas.teufel@kit.edu
>
> **Please ask questions in the forum/discussion board and only contact the Tutor when there are issues with the grading**
---

**Topic**: This exercise sheet will not be graded and serves as an
introduction  to explain the online exercise regulations and to help you to
familiarize yourself with Python, Jupyter and Numpy. The exercises in this
sheet are meant as an appetizer to show you what future exercises could cover.

---

Please add here your group members' names and student IDs. 

Names: 

IDs:

---

# 1.1 AutoGrading

Before jumping straight into the exercise, we first want to provide you with an introduction to the ``auto-grading`` system that is used to grade subsequent exercise notebooks.

Generally, each exercise will be released as a jupyter notebook containing certain tasks related to the lecture topics. The notebooks will contain a mixture of markdown descriptions, already existing code and placeholder were additional code has to be added. After filling in all the required code, the exercise notebooks can be be exported and subsequently uploaded to ILIAS. In the week following the submission deadline, all the submitted notebooks will be automatically graded and the results will be sent to each individual student via email.

## Different Cell Types

For an auto-grading jupyter notebook, there are three types of cells:
- **Read-Only.** Generally all markdown cells and some code cells fall into this category. These cells are annotated with a special "read only" flag and _should not_ be edited under any circumstances. Any changes that are made in these cells will be _discarded_ during the auto-grading process. Read-only code cells can be identified by a special "DO NOT MODIFY" comment at the start.
- **Answers.** These are the only cells that should be edited while working on the exercise. Answer cells usually contain clear comments what kind of code should be added at which points to identify the exercise.
- **Tests.** These are a special case of the *read-only* cell. Test cells contain a series of assert statements that will test the validity and correctness of a previously implemented *answer* cell. For instance, if the goal of a task was to compute a specific value, then the corresponding test cell will check if the variable contains a value of the correct type and within an approximately correct value range, for example.

### Test Assertions

As previously introduced, testing cells contain various assertions which check the validity and the correctness of an implemented answer. Each test cell is associated with a number of *points* &mdash; if a test cell is able to execute *without any exception* then those points are awarded to the final score.

However, the assertions that are visible in the published version of the notebook only constitute the **public tests**. The public tests are a subset of the tests that are supposed to mainly check for the technical validity of the solution (correct data type or shape of tensors) and sometimes offer hints toward the correctness (approximately correct value range). Passing the public tests is a good indication regarding the correctness of the exercise, but *does not guarantee* the points!

In addition to the public tests, each testing cell also contains a set of **hidden tests** which are only applied during grading. Only if both the public *and* the hidden tests are passed, the points of the cell are awarded to the final score.

### BwJupyter

Summer term 2026 is the first time that we are trying out the [BwJupyter Service](https://www.bwjupyter.de/) which is available to every KIT student and lecturer. It provides a centralized opportunity to work on jupyter notebooks, similar to Google Colab.

While completing the notebooks on BwJupyter is the *recommended* way of interacting with the exercise, you can also set up a local jupyter development environment. However, in this case you should make sure to use the exact same dependencies which are also used during grading: https://bwsyncandshare.kit.edu/s/qmPknsTyFs94fyd 

### BwJupyter - AutoGrading Extension

While working on an exercise notebook on bwJupyter, the custom ``aimat-nbgrader`` extension should automatically be installed in your editor (checkbox icon in the sidebar; see image).

<img src="https://bwsyncandshare.kit.edu/s/MX4F8GZjMDeHaP3/download" width="50%">

This extension provides you with some useful information about the progress of the exercise regarding answer and testing cells. In the sidebar panel you can click the ID's to make the editor jump dirrectly to the corresponding cells.

### Additional Remarks

- **Do not delete any cells.** The autograding process relies on the integrity of the cell metadata, which means that cells should not be deleted under any circumstances! If you accidentally delete a cell, then do not just add it again, but re-download the entire notebook instead. Notebooks with broken metadata will fail during grading.
- **Avoid adding additional cells.** Similar problems can occur when additional cells are added to the notebook. Please use the existing answer cells to implement your solutions and avoid adding additional cells, if possible.

## 1.2 Corona

### Preliminaries
If you are not familiar with Python, you may want to learn more about Python
and its basic syntax. Since there are a lot of free and well written tutorials
 online, we refer you to one of the following online tutorials:

* http://www.datacamp.com/community/tutorials/tutorial-jupyter-notebook
* https://www.learnpython.org/
* https://automatetheboringstuff.com/

### COVID-19 Data

To get to know Python's data science workflows, we will analyze some historical data from the
COVID-19 pandemic. Specifically, we will use a dataset on corona infections worldwide from the
European Centre for Disease Prevention and Control (ECDC) covering the year 2020
([link](https://www.ecdc.europa.eu/sites/default/files/documents/COVID-19-geographic-disbtribution-worldwide.xlsx)).

This exercise will introduce you to some of the most commonly used Python libraries for data
analysis and visualization: ``pandas`` for tabular data manipulation, ``numpy`` for numerical
computations, ``matplotlib`` for plotting, and ``scipy`` for curve fitting. We can download the
dataset in Python via the ``requests`` package.

```python
##### DO NOT CHANGE #####
# Standard library
import os
from datetime import datetime

# Data handling
import numpy as np
import pandas as pd
import requests

# Visualization
import matplotlib.dates as mdates
import matplotlib.pyplot as plt

# Curve fitting and machine learning
import scipy.optimize
from sklearn.neural_network import MLPRegressor

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
# Dataset URL and local filename
data_url = "https://www.ecdc.europa.eu/sites/default/files/documents/COVID-19-geographic-disbtribution-worldwide.xlsx"
data_file = "COVID-19-geographic-disbtribution-worldwide.xlsx"

# Download the dataset if it doesn't already exist locally
if not os.path.exists(data_file):
    print("Downloading dataset ...")
    with open(data_file, "wb") as f:
        f.write(requests.get(data_url).content)
    print("Downloading dataset done.")

##### DO NOT CHANGE #####
```

Now, we load the dataset via the data library ``pandas``, which will return a ``DataFrame`` object.
A ``DataFrame`` is the central data structure in pandas — it represents a two-dimensional table
with labeled rows and columns, similar to a spreadsheet or SQL table. We can inspect the first
few rows of the table with the ``.head()`` method:

```python
##### DO NOT CHANGE #####
# Load the Excel file into a pandas DataFrame and display the first rows
corona_world = pd.read_excel(data_file)
corona_world.head()

##### DO NOT CHANGE #####
```

The ``DataFrame`` allows access via its index (row labels) and columns (column names).
Basic Python operators ``[]`` and ``.`` are supported for selecting columns.

**Warning**: Whether a copy or a reference is returned for a setting operation
may depend on the context. This is a common source of bugs when working with pandas.

```python
##### DO NOT CHANGE #####
# Inspect the row index and column names of the DataFrame
print(corona_world.index)
print(corona_world.columns)

##### DO NOT CHANGE #####
```

Now we will select a column, namely ``corona_world["countriesAndTerritories"]``,
and then select the rows where we find ``"Germany"`` via ``corona_countries == "Germany"``.
This works not by specifying row numbers, but by using a boolean array as an index — a
powerful pandas feature called *boolean indexing*.

We can also select multiple columns at once via a list of column names, e.g.
``[["dateRep", "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000"]]``.
The underlying values of a ``DataFrame`` can be obtained as a numpy array via ``.values``
or ``.to_numpy()``.

Note that usually you will interact with a pandas ``DataFrame`` using the accessors
``.iloc`` (integer-based indexing) and ``.loc`` (label-based indexing) for selecting groups
of rows and columns. Both support slicing. For convenience, we flip the final data arrays
so that they start from the past going forward.

```python
##### DO NOT CHANGE #####
# Select only the rows for Germany using boolean indexing
corona_countries = corona_world["countriesAndTerritories"]
corona_germany = corona_world[corona_countries == "Germany"]

# Keep only the columns we need
corona_germany = corona_germany[[
    "dateRep", "cases", "deaths",
    "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000"
]]

# Extract the date and 14-day cumulative cases as numpy arrays
time_germany = corona_germany["dateRep"].values
cc_germany = corona_germany[
    "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000"
].to_numpy()

# Reverse the arrays so they go from earliest to latest date
time_germany = np.flip(time_germany)
cc_germany = np.flip(cc_germany)

##### DO NOT CHANGE #####
```

To visualize the data, we plot the 14-day cumulative number of cases per 100,000 inhabitants
as a function of time for Germany. For this purpose we use the ``matplotlib.pyplot`` plotting
library. You can find a comprehensive user guide with examples
[here](https://matplotlib.org/stable/tutorials/introductory/usage.html#sphx-glr-tutorials-introductory-usage-py).
We will use ``matplotlib`` extensively in the upcoming exercises.

```python
##### DO NOT CHANGE #####
# Set up date formatting for the x-axis
months = mdates.MonthLocator()

# Create a plot of the 14-day cumulative cases over time
fig, ax = plt.subplots()
ax.plot(time_germany, cc_germany)

# Format axes and labels
ax.xaxis.set_major_locator(months)
fig.autofmt_xdate()
plt.grid(True, "both")
plt.title("Germany")
plt.ylabel("14 days cases per 100k")
plt.xlabel("Time [months]")
plt.show()

##### DO NOT CHANGE #####
```

Now it's time for you to answer the following questions. You can work either with
pandas DataFrames directly or convert to numpy arrays and standard Python objects.

Wherever possible, built-in functionality of libraries such as ``numpy`` and ``pandas``
should be preferred over raw Python loops, as these libraries use optimized C/Fortran
implementations under the hood and are significantly faster. Have a look at the documentation
of [numpy](https://numpy.org/doc/) and
[pandas](https://pandas.pydata.org/docs/getting_started/index.html#getting-started)
if you want to know more. We will work with numpy in the following exercises in more detail.

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 1.1 (4 points).</strong> Compute the number of deceased in connection with COVID-19
relative to the population for each country in 2020. Which country has the highest and which has
the lowest mortality rate? Assign the name of the country with the highest mortality rate to the
variable <code>country_max_mortality</code>.

To solve this task you first have to obtain the population per country, for example in a Python
<code>dict</code>. You may want to use the following pandas methods (if you don't know them, look
them up in the <a href="https://pandas.pydata.org/docs/getting_started/index.html#getting-started">pandas documentation</a>):
<code>.unique</code>, <code>.groupby</code>, <code>.mean</code>, <code>.sum</code>,
<code>.sort_values</code>, <code>.get_group</code> and <code>.from_dict</code>.
</div>

<!-- grade_id: ans-1-1-mortality, points: 0, grade: false, solution: true -->
```python
# TASK: Compute the COVID-19 mortality rate (deaths / population) for each country.

# HINT: Use corona_world.groupby("countriesAndTerritories") to group the data by country,
#       then use .sum() for total deaths and .mean() for the population of each country.

# List of all unique country names
countries = []

# Dict mapping country name -> population
population = {}

# Dict mapping country name -> total number of deaths
deaths = {}

# Dict mapping country name -> relative mortality (deaths / population)
rel_deaths = {}

# Name of the country with the highest mortality rate
country_max_mortality = ""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-1-1-mortality, points: 4, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test-1-1-mortality - possible points: 4

# Tests for auto-grading (check variables and types).
assert isinstance(countries, list) or isinstance(countries, np.ndarray)
assert isinstance(rel_deaths, dict)
assert isinstance(population, dict)
assert isinstance(deaths, dict)
assert isinstance(country_max_mortality, str)

# Sanity checks
assert len(countries) > 200, f"Expected more than 200 countries, got {len(countries)}"
assert "Germany" in rel_deaths, "Germany should be in the rel_deaths dictionary"
assert rel_deaths["Germany"] > 0, "Mortality rate for Germany should be positive"
assert country_max_mortality != "", "Please assign the country with highest mortality"
assert country_max_mortality in rel_deaths, f"'{country_max_mortality}' is not in rel_deaths"

# NOTE: The hidden tests will check for the exact number of countries, the exact
#       mortality values for specific countries, and the correct country with the
#       highest mortality rate.

##### DO NOT CHANGE #####
```

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 1.2 (2 points).</strong> Compute and plot the <em>cumulative</em> number of cases and
deaths for Germany. Make sure that the data has the correct chronological order before computing
the cumulative sum. You can either implement your own cumulative sum or simply use
<code>np.cumsum</code>. Store the results in the variables <code>cum_deaths</code> and
<code>cum_cases</code>.
</div>

<!-- grade_id: ans-1-2-cumulative, points: 0, grade: false, solution: true -->
```python
cum_deaths = np.array([])
cum_cases = np.array([])

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-1-2-cumulative, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test-1-2-cumulative - possible points: 2

assert isinstance(cum_deaths, np.ndarray)
assert isinstance(cum_cases, np.ndarray)

# Sanity checks
assert len(cum_deaths) > 300, f"Expected more than 300 data points, got {len(cum_deaths)}"
assert len(cum_cases) > 300, f"Expected more than 300 data points, got {len(cum_cases)}"
assert cum_deaths[-1] > cum_deaths[0], "Cumulative deaths should be increasing overall"
assert cum_cases[-1] > cum_cases[0], "Cumulative cases should be increasing overall"
assert cum_cases[-1] > cum_deaths[-1], "Total cases should exceed total deaths"

# NOTE: The hidden tests will check the exact cumulative values for Germany
#       at specific time points.

##### DO NOT CHANGE #####
```

**Background.** There has been a discussion in Germany whether the "second wave" of infections
in late 2020 was underestimated by many political leaders. The chancellor of Germany, Angela Merkel,
and her team warned the ministers and the public of a dire development at the beginning of October.
In fact, the large increase of infections before Christmas turned out to be very serious, even with
counter-measures in place. In the next tasks, we want to see if we can correctly predict this
development using simple curve fitting and machine learning.

First, we select the data between 01.08.2020 and 01.10.2020 as our training window. Note that we
are using numpy arrays with dtype ``datetime64[ns]``. We can use the ``datetime`` package and its
``datetime`` class for human-readable dates, which can be freely converted into strings, seconds,
or via ``np.datetime64()`` into numpy's datetime format. We select the correct time period and
divide by ``np.timedelta64(1, "D")`` (one day) to obtain an array counting the number of days
since 01.08.2020.

```python
##### DO NOT CHANGE #####
# Define the training window: August 1 to October 1, 2020
date_start = datetime(2020, 8, 1)
date_stop = datetime(2020, 10, 1)

# Create a boolean mask for the selected time range
selection = np.logical_and(
    time_germany >= np.datetime64(date_start),
    time_germany < np.datetime64(date_stop)
)

# Training data: convert dates to "days since start" for curve fitting
x_days = np.array(
    time_germany[selection] - np.datetime64(date_start)
) / np.timedelta64(1, "D")
y_cc_germany = np.array(cc_germany)[selection]

# Test data: everything from the start date onward (for extrapolation)
x_test = time_germany[time_germany >= np.datetime64(date_start)]
x_test_days = (x_test - np.datetime64(date_start)) / np.timedelta64(1, "D")

##### DO NOT CHANGE #####
```

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 1.3 (3 points).</strong> Implement the following analytical functions in Python
using <code>numpy</code> methods such as <code>np.square</code>, <code>np.exp</code>, and
<code>np.power</code>. Each function should accept a numpy array <code>x</code> and the
respective parameters, and return the computed values as a numpy array.

<ul>
<li>$f_1(x) = ax + b$</li>
<li>$f_2(x) = c \, e^{a(x-b)}$</li>
<li>$f_3(x) = ax^2 + bx + c$</li>
<li>$f_4(x) = ax^3 + bx^2 + cx + d$</li>
</ul>
</div>

<!-- grade_id: ans-1-3-functions, points: 0, grade: false, solution: true -->
```python
# Implement functions.
def fun_1(x: np.ndarray, a: float, b: float) -> np.ndarray:
    """Linear function: f(x) = a*x + b."""
    return a * x + b


def fun_2(x: np.ndarray, a: float, b: float, c: float) -> np.ndarray:
    """Exponential function: f(x) = c * exp(a * (x - b))."""
    pass


def fun_3(x: np.ndarray, a: float, b: float, c: float) -> np.ndarray:
    """Quadratic polynomial: f(x) = a*x^2 + b*x + c."""
    pass


def fun_4(x: np.ndarray, a: float, b: float, c: float, d: float) -> np.ndarray:
    """Cubic polynomial: f(x) = a*x^3 + b*x^2 + c*x + d."""
    pass


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-1-3-functions, points: 3, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test-1-3-functions - possible points: 3

assert fun_1(0, 0, 0) == 0
assert fun_2(0, 0, 0, 0) == 0
assert fun_3(0, 0, 0, 0) == 0
assert fun_4(0, 0, 0, 0, 0) == 0

# Correctness checks with simple inputs
assert np.isclose(fun_1(1, 1, 0), 1), "f1(1, a=1, b=0) should equal 1"
assert np.isclose(fun_2(0, 1, 0, 1), 1), "f2(0, a=1, b=0, c=1) should equal 1"
assert np.isclose(fun_3(2, 1, 0, 0), 4), "f3(2, a=1, b=0, c=0) should equal 4"
assert np.isclose(fun_4(1, 1, 1, 1, 1), 4), "f4(1, a=1, b=1, c=1, d=1) should equal 4"

# NOTE: The hidden tests will evaluate each function with specific parameter
#       values and check for the correct numerical output.

##### DO NOT CHANGE #####
```

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 1.4 (4 points).</strong> Use <code>scipy.optimize.curve_fit</code> to fit the
data from 01.08.2020 to 01.10.2020 with each of the four previously implemented functions.
If the fit does not converge, try setting reasonable bounds and initial guesses via the
<code>p0</code> parameter. Plot the fitted functions together with the ground truth data and
visually evaluate the extrapolation quality beyond 01.10.2020. Then answer the analysis questions
in the cell below.
</div>

<!-- grade_id: ans-1-4-curve-fit, points: 0, grade: false, solution: true -->
```python
# Fit f1 (linear) — given as an example
param_1, _ = scipy.optimize.curve_fit(fun_1, x_days, y_cc_germany)

# Fit f2, f3, f4 — use curve_fit with initial guesses where needed.
# Hint: the default p0=[1, 1, 1] works for the polynomials, but for the
# exponential f2 it overflows — curve_fit evaluates exp(a*(x-b)) with
# a=1, x up to ~60, giving exp(59) ≈ 1e25 on the first step.
# Pass a reasonable p0 via the p0 argument: a small positive growth rate
# (a few percent per day), b ≈ 0, and c on the order of y_cc_germany[0].
param_2, _ = None, None  # TODO: fit fun_2
param_3, _ = None, None  # TODO: fit fun_3
param_4, _ = None, None  # TODO: fit fun_4

# YOUR CODE HERE
raise NotImplementedError()

# Evaluate f1 on the test range — given as an example
y_fit_1 = fun_1(x_test_days, *param_1)

# Evaluate f2, f3, f4 on the test range
y_fit_2 = None  # TODO: evaluate fit 2
y_fit_3 = None  # TODO: evaluate fit 3
y_fit_4 = None  # TODO: evaluate fit 4

# YOUR CODE HERE
raise NotImplementedError()

# Plot ground truth and all fitted curves
fig, ax = plt.subplots()
ax.plot(time_germany, cc_germany, label="Ground truth", color="black")
ax.plot(x_test, y_fit_1, label="f1 (linear)")

# YOUR CODE HERE
raise NotImplementedError()

ax.xaxis.set_major_locator(months)
fig.autofmt_xdate()
plt.grid(True, "both")
plt.ylim([-20, 370])
plt.legend()
plt.title("Germany — Curve Fitting Extrapolation")
plt.ylabel("14 days cases per 100k")
plt.xlabel("Time [months]")
plt.show()
```

Answer the following questions by assigning the appropriate values to the given variables:

- Which function extrapolates best to the subsequent development from 01.10.2020 to the end of the year?
- Can you quantify the difference (e.g. via the mean squared error on the test period)?
- Does the answer change if you only fit to the last 14 days (17.09.2020 to 01.10.2020)?
- Which is the best function for 14-day extrapolation, and can you explain why?

<!-- grade_id: ans-1-5-extrapolation, points: 0, grade: false, solution: true -->
```python
# Which function number (1-4) extrapolates best over the full period?
answer_best_fun: int = None

# Does the answer change when fitting only the last 14 days? (True/False)
answer_change_14_days: bool = None

# Which function number (1-4) is best for 14-day extrapolation?
answer_best_fun_14_days: int = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: ans-1-4-explanation, points: 0, grade: false, solution: true -->
```python
# Explain your reasoning in free text.
answer_explanation = ""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-1-4-fit-analysis, points: 3, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test-1-4-fit-analysis - possible points: 3

# Check that all fits were computed
assert param_1 is not None
assert param_2 is not None
assert param_3 is not None
assert param_4 is not None

# Check parameter counts match function signatures
assert len(param_1) == 2, "f1 (linear) should have 2 parameters"
assert len(param_2) == 3, "f2 (exponential) should have 3 parameters"
assert len(param_3) == 3, "f3 (quadratic) should have 3 parameters"
assert len(param_4) == 4, "f4 (cubic) should have 4 parameters"

# Check answer types
assert isinstance(answer_best_fun, int)
assert isinstance(answer_change_14_days, bool)
assert isinstance(answer_best_fun_14_days, int)

# Check answer ranges
assert answer_best_fun in [1, 2, 3, 4], "answer_best_fun must be 1, 2, 3, or 4"
assert answer_best_fun_14_days in [1, 2, 3, 4], "answer_best_fun_14_days must be 1, 2, 3, or 4"

# Internal consistency: answer_change_14_days must agree with whether the two
# best-function numbers differ.
assert answer_change_14_days == (answer_best_fun != answer_best_fun_14_days), \
    "answer_change_14_days must be True iff the best function differs between the full-range and 14-day fits"

# NOTE: The hidden tests will additionally check fitted parameters for correctness
#       and verify the exact answers about the best extrapolating function.

##### DO NOT CHANGE #####
```

<!-- grade_id: test-1-4-explanation, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test-1-4-explanation - possible points: 1

assert isinstance(answer_explanation, str)
assert len(answer_explanation) > 10, "Please provide a meaningful explanation."


##### DO NOT CHANGE #####
```

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 1.5 (2 points).</strong> In addition to analytical curve fitting, we can also use
machine learning to model time series data. The code cell below trains a small neural network
(<code>MLPRegressor</code> from scikit-learn) on 14-day windows to predict the change in
infections for the next day. You do not need to understand this code in detail yet — we will
cover neural networks later in the course.

Your task is to <strong>plot</strong> the neural network predictions (<code>x_nn_test</code>,
<code>y_nn_test</code> and <code>x_nn_val</code>, <code>y_nn_val</code>) together with the ground truth
data and <strong>describe</strong> the problems you observe. How would you improve the model?
Assign your answer as a string to the variable <code>answer_nn</code>.
</div>

```python
##### DO NOT CHANGE #####
# --- Training Data Preparation ---
# Use a sliding window of 14 days as input features.
# The target is the day-over-day change in infection count.
step_size = 14

x_train = np.array([
    y_cc_germany[i : i + step_size]
    for i in range(len(x_days) - step_size)
])

y_train = np.array([
    y_cc_germany[step_size + i] - y_cc_germany[step_size + i - 1]
    for i in range(len(x_days) - step_size)
])

# --- Model Training ---
# Train a simple multi-layer perceptron (MLP) regressor
nn = MLPRegressor(
    hidden_layer_sizes=(100, 100),
    random_state=1,
    max_iter=500,
).fit(x_train, y_train)

# --- Autoregressive Prediction ---
# Starting from the last 14 days of training data, predict 30 days forward
# by feeding each prediction back as input for the next step.
y_nn_test = y_cc_germany[-step_size:]
y_nn_val = y_cc_germany[:step_size]
days_to_predict = 30

for _ in range(days_to_predict):
    # Predict from the end of the training window (extrapolation)
    next_step = nn.predict(np.expand_dims(y_nn_test[-step_size:], axis=0)) + y_nn_test[-1]
    y_nn_test = np.concatenate([y_nn_test, next_step], axis=0)

    # Predict from the start of the training window (validation)
    next_step = nn.predict(np.expand_dims(y_nn_val[-step_size:], axis=0)) + y_nn_val[-1]
    y_nn_val = np.concatenate([y_nn_val, next_step], axis=0)

# --- Time Axes for Plotting ---
x_nn_test = np.arange(-step_size, days_to_predict) * np.timedelta64(1, "D") + np.datetime64(date_stop)
x_nn_val = np.arange(0, days_to_predict + step_size) * np.timedelta64(1, "D") + np.datetime64(date_start)

##### DO NOT CHANGE #####
```

You can play around with the hyper-parameters of the neural network (parameters not optimized
during training but used to control the learning process, such as ``hidden_layer_sizes``,
``max_iter``, or ``random_state``) and see how the prediction changes.

<!-- grade_id: ans-1-6-neural-net, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: ans-1-5-explanation, points: 0, grade: false, solution: true -->
```python
# Describe the problems you observe and how you would improve the model.
answer_nn = ""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-1-5-neural-net, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test-1-5-neural-net - possible points: 2

assert isinstance(answer_nn, str)
assert len(answer_nn) > 10, "Please provide a meaningful explanation."

# NOTE: The hidden tests will check your explanation using AI grading.

##### DO NOT CHANGE #####
```

# Submitting your solution

### Verify Assignment

**This semester, the Verify Assignment feature is only available on bwJupyter.** If you are working locally, you can still run the test cells manually to check your solutions.

Before submitting, use the **Verify Assignment** button in the toolbar to run a comprehensive check of your
notebook. This will execute all cells, check for placeholder code, and show you which tests pass
or fail along with your current point total.

Ideally, the verification should pass without any errors. It serves as a general sanity check for
your notebook based on the non-hidden assert statements. Note that passing the verification does
not automatically mean that you will receive full points — there are additional hidden tests that
run during grading.

### Upload to ILIAS

To upload your notebook first go to ``File > Download`` to get the ``.ipynb`` file.

