# Exercise Sheet No. 5

---

> Machine Learning for Natural Sciences, Summer 2026, Prof. Pascal Friederich
>
> Tutor: houssam.metni@kit.edu
>
> **Please ask questions in the forum/discussion board and only contact the tutor when there are issues with the grading (after you have received your grades).**

---

**Topic**: This exercise sheet will focus on feed-forward neural networks, their implementation and training, as well as an application to materials science.

Please add here your group members' names and student IDs. 

You are encouraged to work in groups of a maximum of 3 people, however **each of you** has to submit a solution.

Names: 

IDs:

# Application

You are learning about applications of NNs in materials science. To give you a little insight, we will use an application in materials science here:

# Organic Solar Cells
For organic materials to become semi-conducting, electrons must be delocalized in the molecule. For electrons to be delocalized, a high level of conjugation is necessary.  
When single and double bonds are alternating in an organic molecule, electrons can move. When we think about an aromatic ring, like benzene, it is not defined where the double bonds would form, so they can move around the ring and are delocalized along the whole aromatic system.

```python
# Installing RDKit
import sys
import subprocess

try:
    import rdkit
    print("RDKit ready!")
except ImportError:
    try:
        subprocess.check_call(["uv", "pip", "install", "rdkit"])
        print("RDKit installed via uv!")
    except FileNotFoundError:
        subprocess.check_call([sys.executable, "-m", "pip", "install", "rdkit"])
        print("RDKit installed via pip!")
```

```python
##### DO NOT CHANGE #####
from rdkit import Chem
from rdkit.Chem import Draw
from rdkit.Chem.Draw import IPythonConsole

mol = Chem.MolFromSmiles('c1ccccc1')
mol

##### DO NOT CHANGE #####
```

The electrons that have higher energies are part of what is called the highest occupied molecular orbit (HOMO). This is basically the equivalent of the valence band in classical semiconductors.  

The equivalent of the conduction band is the lowest level at higher energies that is unoccupied or the lowest unoccupied molecular orbit (LUMO). The gap between those two levels is the bandgap of organic semiconductors.

For organic photovoltaic cells this bandgap needs to be small enough so that visible light can excite an electron from the HOMO to the LUMO. This requires a high level of conjugation and hence aromatic systems (Figure 1)

<a title="Alevina89, CC BY-SA 3.0 &lt;https://creativecommons.org/licenses/by-sa/3.0&gt;, via Wikimedia Commons" href="https://commons.wikimedia.org/wiki/File:Homo-lumo_gap.tif"><img width="512" alt="Homo-lumo gap" src="https://upload.wikimedia.org/wikipedia/commons/thumb/5/58/Homo-lumo_gap.tif/lossless-page1-607px-Homo-lumo_gap.tif.png"></a>
<p style="text-align:center;font-size:80%;font-style:italic">
    Figure 1: Conjugation induced LUMO reduction.
</p>

For the development of organic semiconductors, HOMO and LUMO can be simulated by density functional theory (DFT) using different levels of theory. All you need to know now, is that depending on the level of theory  - or better the details of the approximations taken - the calculation of properties like HOMO and LUMO from the molecule can take hours.:

This is a problem for high-throughput screening if one wants to discover new materials. It is extremely costly to evaluate e.g. 100,000 molecules for their properties with methods that are precise enough. Hence, one usually tries to do a detailed simulation only on a subset of molecules and then train a ML-model to predict the properties of interest on the labeled data.

## Dataset
The dataset contains the simulated LUMO and HOMO values for 51,247 organic molecules. Additionally, it contains 63 molecule descriptors that were calculated using [rdkit](https://www.rdkit.org/docs/index.html)

```python
##### DO NOT CHANGE #####
# get dataset
import pandas as pd
import numpy as np
from matplotlib import pyplot as plt
import requests

url = 'https://bwsyncandshare.kit.edu/s/WcG2c96Xq82gB4E/download'
response: str = requests.get(url).content.decode('utf-8')
with open('OPV.csv','w') as file:
    file.write(response)

df = pd.read_csv('OPV.csv', header=1)
df.describe()

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
# HOMO/LUMO visualization
df['LUMO'].hist(bins=100, label='LUMO')
df['HOMO'].hist(bins=100, label='HOMO')
plt.legend()
plt.show()

##### DO NOT CHANGE #####
```

# Linear regression benchmark
For a first shot we can train a ridge regression to predict the `labels` from the `mol_descriptors`.

This time we will use `sklearn`. Additionally we will also pre-process the features with a standard scaler to shift them to zero mean and scale them to unit variance.

As already mentioned, the calculation of DFT properties can be very costly. Hence, we want to have a model that extrapolates well to unseen data. Hence, we will use only 20% of the dataset for training and test on 80%.

```python
##### DO NOT CHANGE #####
from sklearn.linear_model import LinearRegression, Ridge
from sklearn.preprocessing import StandardScaler
from sklearn.model_selection import train_test_split
from sklearn.metrics import r2_score

mol_descriptors = df.columns[2:]

X = df[mol_descriptors].values
y = df[['LUMO','HOMO']].values

##### DO NOT CHANGE #####
```

Next we use the [`train_test_split()`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html?highlight=train_test_split#sklearn.model_selection.train_test_split) to split of 20% of `X` and `y` as train set and use the rest as test set:

```python
##### DO NOT CHANGE #####
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.8, random_state=0)

##### DO NOT CHANGE #####
```

Next, we fit a `StandardScaler` object to scale `X_train`, thus obtaining `X_train_scaled`, and use the parameters of the scaler to transform `X_test` to `X_test_scaled` (so using the same scaler from the train test). 

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.1 (1 point).</strong> Fit a <code>StandardScaler</code> to <code>X_train</code>, assign it to <code>scaler</code>, and produce <code>X_train_scaled</code> and <code>X_test_scaled</code> using the same fitted scaler (do not refit on the test data).
</div>

<!-- grade_id: Scaling_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.1 : Assign a StandardScaler object to scaler and obtain the scaled features as X_train_scaled and X_test_scaled
scaler = None
X_train_scaled = None
X_test_scaled = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: Scaling_Test, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: Scaling_Test - possible points: 1

# Scaler - 1 point

assert isinstance(scaler, StandardScaler), "The scaler should be an instance of the sklearn StandardScaler"
assert X_train_scaled.shape == X_train.shape, "The shape of the scaled array should be the same as the unscaled"
assert X_test_scaled.shape == X_test.shape, "The shape of the scaled array should be the same as the unscaled"

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
# renaming variables for simplicity

X_train = X_train_scaled
X_test = X_test_scaled

##### DO NOT CHANGE #####
```

Now use the [`Ridge()`](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.Ridge.html?highlight=ridge#sklearn.linear_model.Ridge) model to fit a ridge regression with standard parameters to the train data, and assign the predictions of the fitted model on the test data to `y_pred`:

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.2 (1 point).</strong> Instantiate a <code>Ridge()</code> regression model, fit it on the (scaled) training set, and assign its predictions on the test set to <code>y_pred</code>.
</div>

<!-- grade_id: Ridge_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.2 : Instantiate a Ridge() model as model, fit it on the train set and assign the predictions to y_pred
model = None
y_pred = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
##### DO NOT CHANGE #####
r2_lumo_ridge = r2_score(y_test[:, 0], y_pred[:, 0])
r2_homo_ridge = r2_score(y_test[:, 1], y_pred[:, 1])
print(f'R2 LUMO: {r2_lumo_ridge}\nR2 HOMO: {r2_homo_ridge}')

##### DO NOT CHANGE #####
```

<!-- grade_id: Ridge_Test, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: Ridge_Test - possible points: 1

# Ridge Regression - 1 point

assert y_pred.shape[0] == y_test.shape[0], 'assert that the shapes are the same'
assert y_pred.shape[1] == y_test.shape[1], 'assert that the shapes are the same'
assert r2_lumo_ridge > 0.70, 'the r2 score should be of at least around 0.7'
assert r2_homo_ridge > 0.70, 'the r2 score should be of at least around 0.7'
# Possible hidden tests

##### DO NOT CHANGE #####
```

Now we can plot the correlations of the the true and predicted values:

```python
##### DO NOT CHANGE #####
fig, axs = plt.subplots(1, 2, figsize=(12, 6))

axs[0].scatter(y_test[:, 0], y_pred[:, 0], alpha=0.2, s=5)
axs[0].plot([-6, -1], [-6, -1], 'k')
axs[0].set_title(f'R2 LUMO: {r2_lumo_ridge}')
axs[0].set_xlabel('true LUMO')
axs[0].set_ylabel('predicted LUMO')
axs[0].set_xlim([-6, -1])
axs[0].set_ylim([-6, -1])

axs[1].scatter(y_test[:, 1], y_pred[:, 1], alpha=0.2, s=5)
axs[1].plot([-9, -3], [-9, -3], 'k')
axs[1].set_title(f'R2 HOMO: {r2_homo_ridge}')
axs[1].set_xlabel('true HOMO')
axs[1].set_ylabel('predicted HOMO')
axs[1].set_xlim([-9, -3])
axs[1].set_ylim([-9, -3])

plt.show()

##### DO NOT CHANGE #####
```

As you can see, the ridge regression learns something, but the predictions could be better...

Hence, we will apply a non-linear model, namely a:

# Neural Network

## Inspiration from Neuroscience
You saw that feedforward neural networks can be thought of as 'function approximation machines' [Goodfellow et al. 2016](https://www.deeplearningbook.org/contents/mlp.html).
Some ideas of today's artificial neural networks are loosely inspired by neuroscientific models of biological networks.
For a different perspective we can try to make a connection to models of biological neural networks (while remembering the goal of today's neural networks is not to get a model of the brain but rather to approximate some function).

The fundamental units of the brain are neurons which can also be thought of as information messengers. Neurons receive, transmit and transform information in the form of electrical impulses.
A neuron integrates information from its upstream neurons at the dendrites, creating a post-synaptic potential. 
At the axon hillock this potential is encoded into a spike train of action potentials. These are sent down the axon until they reach a synapse of an axon terminal. At the synapse the spike train is again decoded into a pre-synaptic potential and depending on the incoming spike frequency more or less neurotransmitters are released to pass the signal on to the next neuron (Figure 2).
<p style="text-align:center;font-size:80%;font-style:italic">
    <img src="https://images.topperlearning.com/topper/tinymce/imagemanager/files/Synapse_between_Neurons.jpg", width="70%">
    <br>
    Figure 2: Simplified sketch of a neuron.
</p>


## Forward Pass

For today's artificial neural networks this was simplified in the following way:  
Instead of simulating spikes, we only simulate a spike frequency as a continuous value. Depending on the post-synaptic potential, this frequency can be higher or lower. E.g. neurons do not respond at all until a specific post-synaptic potential is reached and also have a maximum spike-frequency. This implies two things:  
**1. The transformation from post-synaptic potential to spike-frequency is non-linear. (activation function)**  
**2. Each neuron has a different base activity. (bias)**

In the following mathematical definitions, lower-case letters denote scalars, while upper-case letters denote vectors or matrices. No dot or $\cdot$ denotes a dot product of vectors or matrices.

### Single neuron
Given a subjected neuron $l$ receives inputs from an upstream layer $k$ and the upstream layer consists of three neurons with the outputs $x_1, x_2, x_3$ via connections with weights $\theta_{1,l}, \theta_{2,l}, \theta_{3,l}$, we can compute the weighted sum as state $h_l(X)$ of the subjected neuron by linear algebra as the product of the row vector $X_k$ and the column vector $\Theta_{k,l}$. Additionally, we add a bias $b$ to the neuron.

\begin{align}
h_l(X) &= X_k \cdot \Theta_{k,l} + b_l\\
&=
\begin{bmatrix}
    x_1 & x_2 & x_3
\end{bmatrix} 
\cdot
\begin{bmatrix}
    \theta_{1,l} \\ \theta_{2,l} \\ \theta_{3,l}
\end{bmatrix}
+ b_l\\
&=x_1 \theta_{1,l} +  x_2 \theta_{2,l} + x_3 \theta_{3,l} + b_l
\end{align}

And the output or activation $a_l(x)$ of the subjected neuron in terms of spike-frequency is computed using a non-linear activation function $\sigma()$:

\begin{align}
a_l(X) &= \sigma(h_l(X))
\end{align}

### Neuron Layer
Now given that the subjected $l$ is not only a single neuron but e.g. a layer of two neurons, the notation is the same with the only difference that we add a column to the weight column vector $\Theta_{k,l}$ changing it to a $3 \times 2$ dimensional matrix to calculate the $1 \times 2$ dimensional row state vector $H_l(X)$:

\begin{align}
    H_l(X) &= X_k \cdot \Theta_{k,l} + B_l\\
    &=
    \begin{bmatrix}
        x_1 & x_2 & x_3
    \end{bmatrix} 
    \cdot
    \begin{bmatrix}
        \theta_{1,1} & \theta_{1,2} \\
        \theta_{2,1} & \theta_{2,2} \\
        \theta_{3,1} & \theta_{3,2} 
    \end{bmatrix}
    +
    \begin{bmatrix}
        b_1 & b_2 
    \end{bmatrix} \\
    &=
    \begin{bmatrix}
        \theta_{1,1} x_1 + \theta_{1,2} x_2 + \theta_{1,3} x_3 + b_1 &
        \theta_{2,1} x_1 + \theta_{2,2} x_2 + \theta_{2,3} x_3 + b_2
    \end{bmatrix}
\end{align}

And the activation of the layer:

\begin{align}
A_i(X) &= \sigma(H_i(X))
\end{align}

For the bias there exist two approaches:
1. One can see the bias as an additional input with a constant $1$ from the previous layer and a learnable weight vector.
2. Adding the bias simply as row vector that added to the neuron state $H(x)$.

Here we use approach 2.

### Batches
Usually data is processed in batches, so not only one single sample is processed at a time but multiple samples as this is much faster when calculated using vectorization. In our case, samples are rows with the length of the features. So when we pass 10 samples of our features (the `mol_descriptors`), we pass a matrix of shape $10 \times 63$. This already fits well with our current notation and it will return a $10 \times 32$ state matrix.


### Weight initialization
To compute anything, the weights must be different from $0$. There are several approaches to initialize weights. We will use the approach of [Glorot et Al., 2010](http://proceedings.mlr.press/v9/glorot10a.html).
The initial weights are drawn form a uniform distribution $U(-z, z)$ with the limit $z$ calculated as:
\begin{align}
    z &= \sqrt{\frac{6}{k+l}}
\end{align}
...with input length $k$ and output length $l$.


Just to get an idea how this all works in python we will build a layer with 32 neurons and feed all `mol_descriptors`.  
Initialize the weight matrix `theta` using numpy functions:

```python
##### DO NOT CHANGE #####
sample = X_train[0:10]
n_inputs = sample.shape[1]
n_outputs = 32

##### DO NOT CHANGE #####
```

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.3 (1 point).</strong> Initialise the weight matrix <code>theta</code> using the Glorot-uniform scheme described above, with shape <code>(n_inputs, n_outputs)</code>.
</div>

<!-- grade_id: Weights_Init_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.3 : Initialize the weights theta after Glorot et Al. 2010 (glorot uniform):
theta = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: Weights_Init_Test, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: Weights_Init_Test - possible points: 1

# Weights - 1 point

assert theta.shape[0] == n_inputs, "Your weight shape doesn't match!"
assert theta.shape[1] == n_outputs, "Your weight shape doesn't match!"

# Hidden asserts for the limits of the uniform distribution

##### DO NOT CHANGE #####
```

Now calculate h(x) for the sample using your weights:

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.4 (2 points).</strong> First initialise the bias row vector <code>b</code> as zeros (length <code>n_outputs</code>), then compute the layer states <code>h(x) = sample · theta + b</code> and assign the result to <code>h</code>. (Two solution cells follow — one for each part.)
</div>

<!-- grade_id: Bias_Init_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.4 : Initialize the bias row vector with zeros:
b = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: hx_Question, points: 0, grade: false, solution: true -->
```python
# Calculate h(x) using sample, theta and b
# The bias won't change anything but it will check for correct shapes.
h = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: hx_Test, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: hx_Test - possible points: 2

# States - 2 points

assert h.shape[0] == sample.shape[0]
assert h.shape[1] == n_outputs
# Possible hidden asserts

##### DO NOT CHANGE #####
```

## Activation functions

Our model will be a regression model. For the output of the model we need to fit the unscaled `y`. Hence we need an unbound linear activation function for the output:

\begin{align}
    linear(H(x)) &= H(x)
\end{align}

As non-linear activation function for the hidden layers we will use the ReLu function:

\begin{align}
    relu(H(x)) &= 
\begin{cases} 
    0~\text{ if }~h_l(x) \leq 0\\
    h_l(x)~\text{ if }~h_l(x)>0
\end{cases}
\end{align}

The ReLu has the biological motivation of having a minimal threshold until it starts outputting values different from 0. Additionally, it yields sparse activations in the network, which is usually favorable, and its gradient is trivial to calculate.

Please calculate the ReLu of your previously calculated `h`:

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.5 (1 point).</strong> Compute the ReLu activation <code>a</code> of the layer states <code>h</code> from the previous task.
</div>

<!-- grade_id: relu_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.5 : Calculate the activation a by computing the ReLu of h
a = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: relu_Test, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: relu_Test - possible points: 1

# ReLu - 1 point

assert np.min(a) >= 0
# Possible hidden asserts

##### DO NOT CHANGE #####
```

## Backpropagation

### Weights
From the lecture recall the following chain to get the gradient of the error $J$ with respect to the weights $\Theta$ of the output layer:

\begin{align}
\frac{\partial J}{\partial \Theta} &= \underset{(1)}{\frac{\partial J}{\partial a}} 
                                      \underset{(2)}{\frac{\partial a}{\partial h}}
                                      \underset{(3)}{\frac{\partial h}{\partial \Theta}} \\
\end{align}

For our case the error is the MSE defined as:
\begin{align}
J &= \frac{1}{2m} \sum_{i=1}^m \left( a^{(i)} - y^{(i)} \right)^2
\end{align}
...for $m$ samples $i$.

Everything else is given. Now calculate the partial derivatives for a single sample $i$ of:  
**(1) the MSE  
(2) the linear and ReLu functions  
(3) $h(x)$**

```python
##### DO NOT CHANGE #####
from IPython.display import display, Markdown

display(Markdown(r"\begin{align}"
                 r"\frac{\partial \text{MSE}}{\partial a^{(i)}} &= \frac{1}{m} \left( a^{(i)} - y^{(i)} \right) \tag{1}\\"
                 r"\frac{\partial \text{ReLu}^{(i)}}{\partial h^{(i)}} &= "
                 r"\begin{cases} "
                 r"0~\text{ if }~h^{(i)}(x) \leq 0\\"
                 r"1~\text{ if }~h^{(i)}(x)>0"
                 r"\end{cases} \tag{2}\\"
                 r"\frac{\partial \text{linear}^{(i)}}{\partial h^{(i)}} &= 1 \tag{3}\\"
                 r"\frac{\partial h^{(i)}}{\partial \theta^{(i)}} &= x^{(i)} \tag{4}"
                 r"\end{align}"))

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
### in case the above cell does not render, this is what the equations should look like ###
import requests
import matplotlib.image as mpimg
url = 'https://bwsyncandshare.kit.edu/s/Tc8T2DjNaqqFENw/download'
response: str = requests.get(url).content
with open('equations.png','wb') as img:
    img.write(response)
img = mpimg.imread('equations.png')

# Display the image
plt.imshow(img)
plt.axis('off')  # Hide the axes
plt.show()

##### DO NOT CHANGE #####
```

Now let's take the above calculated `a` as output and calculate some random true labels `y_` for testing:

```python
##### DO NOT CHANGE #####
y_ = np.random.uniform(low=0, high=3, size=[10, 32])

##### DO NOT CHANGE #####
```

Next we can use the derivatives from above to calculate the gradient for one layer:

\begin{align}
\frac{\partial J}{\partial \Theta} &= \underset{(1)}{\frac{\partial J}{\partial a}} 
                                      \underset{(2)}{\frac{\partial a}{\partial h}}
                                      \underset{(3)}{\frac{\partial h}{\partial \Theta}} \\
\end{align}

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.6 (3 points: 1 + 1 + 1 for the three derivatives).</strong> Implement the three partial derivatives needed for one layer's backpropagation: <code>dh_dtheta</code>, <code>da_dh</code> (using the ReLu derivative), and <code>dJ_da</code> (using the MSE rule with batch size <code>m = 10</code>).
</div>

<!-- grade_id: gradient_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.6
# We start from the back:
# (3) is trivial
dh_dtheta = None
# YOUR CODE HERE
raise NotImplementedError()

# (2) for ReLu: Here you need a reference to h and maybe two steps
da_dh = None
# YOUR CODE HERE
raise NotImplementedError()

# (1) m is the batch size!
dJ_da = None
# YOUR CODE HERE
raise NotImplementedError()

dJ_dTheta = np.dot(dh_dtheta.T, (dJ_da * da_dh)) #implementation to match expectations of numpy in terms of shapes

# The error gradient with respect to the weights and the shape of the weights should agree:
print(dJ_dTheta.shape, theta.shape)
```

<!-- grade_id: gradient_Tests1, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: gradient_Tests1 - possible points: 1

# dh dtheta - 1 point

assert dh_dtheta.shape[0] == 10
assert dh_dtheta.shape[1] == 63
# Hidden test for the content of dh_dtheta

##### DO NOT CHANGE #####
```

<!-- grade_id: gradient_Tests2, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: gradient_Tests2 - possible points: 1

# da dh - 1 point

assert da_dh.shape[0] == 10
assert da_dh.shape[1] == 32
# Hidden test for the content of drelu_dh

##### DO NOT CHANGE #####
```

<!-- grade_id: gradient_Tests3, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: gradient_Tests3 - possible points: 1

# dJ dh - 1 point

assert dJ_da.shape[0] == 10
assert dJ_da.shape[1] == 32
# Hidden test for the content of dJ_da


##### DO NOT CHANGE #####
```

We can then update the weights for the next step using:
\begin{align}
    \Theta_{t+1} &= \Theta_t - \alpha \cdot \frac{\partial J}{\partial \Theta} \\
\end{align}
...with learning rate $\alpha$.

### Bias 

Besides the weights we also need to fit the bias. The bias can be derived in the same manner, only that instead of $x$ the input is $1$. We simply multiply $1$ with the bias weights. Therefore (3) only for the bias collapses to:
\begin{align}
\frac{\partial h^{(i)}_b}{\partial \theta^{(i)}_b} &= x^{(i)}_b = 1\tag{3}
\end{align}

Next we will build a simple feed forward network.

For backpropagation we are still missing the one element connecting the layers, namely $\frac{\partial a^2}{\partial a^1}$:
\begin{align}
a^2 &= \Theta^2 a^1 \\
\frac{\partial a^2}{\partial a^1} &= \Theta^2
\end{align}

With this and the above derived equations we can calculated the gradient of the first (hidden) layer weights regarding the output error as:
\begin{align}
\frac{\partial J}{\partial \Theta^1} &= \underbrace{\frac{\partial J}{\partial a^2}
                                                    \frac{\partial a^2}{\partial a^1}}_\text{(a.)}
                                        \underbrace{\frac{\partial a^1}{\partial h^1}
                                                    \frac{\partial h^1}{\partial \Theta^1}}_\text{(b.)} \\
\end{align}
**The part (a.) can be calculated in the second (output) layer and is returned as upstream gradient to continue the backpropagation process.**

**In the first (hidden) layer we can then use this gradient and combine it with part (b.) to obtain the gradient for the weights of layer 1 with respect to the output error.**

**Further clarifications : upstream/downstream gradients**

${\frac{\partial J}{\partial a^2}}$ is the gradient of the loss $J$ with respect to the output $a^2$ and is also an upstream gradient received directly from the loss function.

${\frac{\partial a^2}{\partial a^1}}$ is the local gradient of the second layer with respect to the first layer's activation

The product of both denoted $(a.)$ represents the downstream gradient asked for below in the `OutputLayer` class.

The combination of the upstream gradient from the loss function and the local gradient of the second layer's output with respect to the first layer's activation is computed at the second layer. This combined gradient (downstream gradient) is then passed back to the first layer as the upstream gradient to continue the backpropagation process.

Let's put this into work by building a feed forward network for our regression task with:

1. The 63 `mol_descriptors` as input $x^{(i)}$
2. One hidden layer $1$ with 32 neurons and a ReLu activation: `class HiddenLayer`
3. One output layer $2$ with two outputs (LUMO and HOMO) with a linear (no) activation function: `class OutputLayer`

Both have a forward pass, which calculates the output of the layer, a backward pass which calculates the gradients and an update function that updates the weights using the gradients and the learning rate.

Only the `OutputLayer` has to return the part (a.) from above to then pass it on to the backward pass of the `HiddenLayer`.

We start implementing from the bottom up with the `OutputLayer` (layer 2). Keep in mind that you should use `n_inputs` and `n_outputs` and not specific numbers (63, 32 or so) when you define things inside of your class. 

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.7 (5 points: 1 forward + 2 backward + 2 update).</strong> Implement the <code>OutputLayer</code> class — its <code>forward</code>, <code>backward</code>, and <code>update</code> logic — using the math from the section above. The class skeleton with TODO markers is provided.
</div>

<!-- grade_id: outputlayer_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.7
class OutputLayer:

    def __init__(self, n_inputs: int, n_outputs: int):
        # Initialize (n_inputs X n_outputs)-dimensional weight matrix self.theta with a
        # Glorot et Al. 2010 uniform initialization:
        self.theta = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Initialize the bias vector self.b to zeros:
        self.b = None
        # YOUR CODE HERE
        raise NotImplementedError()

    def forward(self, input_vector):
        self.input = input_vector
        # Compute the states h(x) as self.h
        self.h = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # As this is a linear layer a(x) = h(x)
        self.a = self.h
        return self.a

    def backward(self, y_predicted, y_true):
        # HINT: as we do things backwards you might have to transpose some matrices

        # partial derivative of the states with respect to the weights
        dh2_dtheta2 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # partial derivative of activations with respect to the states
        # One as we have a linear/no activation for the regression output
        da2_dh2 = 1

        # partial derivative of the error (MSE) with respect to the acivation
        # infer the batch size from the input shape for normalization
        dJ_da2 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Gradient of the weights with respect to the error
        # for the weight updates for this layer:
        self.dJ_dTheta2 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Gradient of the bias with respect to the error
        # Recall using dh2_db2 = 1 and handle the batch size by summation
        self.dJ_db2 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # The downstream gradient for layer 1.
        # employing da2_da1 = theta
        downstream_gradient = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return downstream_gradient

    def update(self, learning_rate):
        # You don't need to change this
        # HINT:
        # If your model gets worse instead of better make sure you calculate the correct MSE
        self.theta = self.theta - learning_rate * self.dJ_dTheta2
        self.b = self.b - learning_rate * self.dJ_db2

```

```python
##### DO NOT CHANGE #####
X_sample = X_train[0:100, :]
y_sample = y_train[0:100, :]

##### DO NOT CHANGE #####
```

<!-- grade_id: outputlayer_Test1, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: outputlayer_Test1 - possible points: 1

# Output Forward Pass - 1 point

l2 = OutputLayer(X_sample.shape[1], y_sample.shape[1])
y_pred_sample = l2.forward(X_sample)

assert y_sample.shape[0] == y_pred_sample.shape[0]
assert y_sample.shape[1] == y_pred_sample.shape[1]

##### DO NOT CHANGE #####
```

<!-- grade_id: outputlayer_Test2, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: outputlayer_Test2 - possible points: 2

# Output Backward Pass - 2 points

downstream_gradient = l2.backward(y_pred_sample, y_sample)

assert downstream_gradient.shape[0] == X_sample.shape[0]
assert downstream_gradient.shape[1] == X_sample.shape[1]

##### DO NOT CHANGE #####
```

<!-- grade_id: outputlayer_Test3, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: outputlayer_Test3 - possible points: 2

# Output Weight Update - 2 points

l2 = OutputLayer(X_sample.shape[1], y_sample.shape[1])
y_pred_before = l2.forward(X_sample)

for i in range(100):
    y_pred_sample = l2.forward(X_sample)
    l2.backward(y_pred_sample, y_sample)
    l2.update(0.05)

y_pred_after = l2.forward(X_sample)
r2_before = r2_score(y_sample, y_pred_before)
r2_after = r2_score(y_sample, y_pred_after)

assert r2_before < r2_after

##### DO NOT CHANGE #####
```

Next is the `HiddenLayer`. There are mainly two main differences you have to keep in mind here:
1. You have to work with the ReLu activation, so $\frac{\partial a}{\partial h}$ isn't $1$ anymore.
2. We use the `upstream_gradient` for the backward pass, provided by the `backward` method of the `OutputLayer`.

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.8 (3 points: 1 for the forward pass + 2 for the weight update).</strong> Implement the <code>HiddenLayer</code> class. The key differences from <code>OutputLayer</code> are (a) the ReLu activation in the forward pass and (b) using the upstream gradient passed into <code>backward</code> instead of computing it from <code>y_true</code> and <code>y_pred</code>. The class skeleton with TODO markers is provided.
</div>

<!-- grade_id: hiddenlayer_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.8
class HiddenLayer:

    def __init__(self, n_inputs: int, n_outputs: int):
        # Initialize the weight matrix self.theta with a
        # Glorot et Al. 2010 uniform initialization:
        self.theta = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Initialize the bias vector to zeros:
        self.b = None
        # YOUR CODE HERE
        raise NotImplementedError()

    def forward(self, input_vector):
        self.input = input_vector
        # Compute the states h(x) as self.h
        self.h = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Compute the activations a(x) as self.a with ReLu activation
        self.a = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return self.a

    def backward(self, upstream_gradient):
        # HINT: as we do things backwards you might have to transpose some matrices

        # Gradient of the states with respect to the inputs (trivial):
        dh1_dtheta1 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Gradient of the activations with respect to the states:
        # Remember to apply ReLu here
        da1_dh1 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Gradient of the error with respect to the weights:
        # Now we can finally use the upstream gradient...
        self.dJ_dTheta1 = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # And for the bias, similar to the output layer but this time
        # using the upstream gradient:
        self.dJ_db1 = None
        # YOUR CODE HERE
        raise NotImplementedError()

    def update(self, learning_rate):
        # You don't need to change this
        self.theta = self.theta - learning_rate * self.dJ_dTheta1
        self.b = self.b - learning_rate * self.dJ_db1

```

```python
##### DO NOT CHANGE #####
# We create a sample from the existing data that matches the output dimension 
# and also shift it above zero, so it is learnable with ReLu:

n_hidden = 32

X_sample = X_train[0:100, :]
y_sample = np.array([y_train[i * 100:i * 100 + 100, 0] for i in range(n_hidden)]).T
y_sample = y_sample - y_sample.min()


##### DO NOT CHANGE #####
```

<!-- grade_id: hiddenlayer_Test1, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: hiddenlayer_Test1 - possible points: 1

# Hidden Forward Pass - 1 point

l1 = HiddenLayer(X_sample.shape[1], n_hidden)

y_pred_sample = l1.forward(X_sample)
assert y_pred_sample.shape[0] == X_sample.shape[0]
assert y_pred_sample.shape[1] == n_hidden
assert y_pred_sample.min() >= 0

##### DO NOT CHANGE #####
```

<!-- grade_id: hiddenlayer_Test2, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: hiddenlayer_Test2 - possible points: 2

# Hidden Weight Update - 2 points

l1 = HiddenLayer(X_sample.shape[1], n_hidden)
y_pred_before = l1.forward(X_sample)

for i in range(100):
    y_pred_sample = l1.forward(X_sample)
    downstream_gradient = 1 / 100 * (y_pred_sample - y_sample)
    l1.backward(downstream_gradient)
    l1.update(0.05)

y_pred_after = l1.forward(X_sample)
r2_before = r2_score(y_sample, y_pred_before)
r2_after = r2_score(y_sample, y_pred_after)

assert r2_before < r2_after

##### DO NOT CHANGE #####
```

## NN Training

Okay now let's see whether we can be better with our network than the benchmark ridge regression.

As said before we will train our network in batches (`batch_size`) and for `n_epochs`.  
Per epoch we therefore have to pass `X_train.shape[0] // batch_size` batches.  
For each epoch we randomly shuffle the dataset. Using `np.random.permutation(X_train.shape[0])` we generate a randomly shuffled index. Indexing X and y with slices of this shuffled index, creates differently shuffled batches for each epoch.

```python
##### DO NOT CHANGE #####
n_hidden = 64
lr = 0.05
n_epochs = 60
batch_size = 100
n_batches = X_train.shape[0] // batch_size

l1 = HiddenLayer(X_train.shape[1], n_hidden)
l2 = OutputLayer(n_hidden, y_train.shape[1])

y_pred_before = l2.forward(l1.forward(X_test))

##### DO NOT CHANGE #####
```

<div style="border: 1px solid #CEB037; border-radius: 3px; padding: 6px; background-color: #faf7e0ff; color: black;">
<strong>🛠️ Task 5.9 (3 points).</strong> Implement the inner body of the training loop: run the forward pass through both layers, then run the backward pass on both layers, and finally call their <code>update</code> methods. A working implementation should improve test R² over the random initialisation and beat the ridge baseline.
</div>

<!-- grade_id: training_Question, points: 0, grade: false, solution: true -->
```python
# Task 5.9
for epoch in range(n_epochs):
    permutation = np.random.permutation(X_train.shape[0])

    for batch in range(n_batches):
        start = batch * batch_size
        end = start + batch_size
        X_batch = X_train[permutation[start:end]]
        y_batch = y_train[permutation[start:end]]

        # Get the predictions
        y_pred = None
        # YOUR CODE HERE
        raise NotImplementedError()

        # Do the backward pass of both layers and update the weights
        # YOUR CODE HERE
        raise NotImplementedError()

    y_pred = l2.forward(l1.forward(X_test))
    if epoch % 4 == 0:
        print(f"Epoch {epoch}: Test R2 = {r2_score(y_test, y_pred)}")

```

<!-- grade_id: training_Test, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: training_Test - possible points: 2

# Full Training - 2 points

y_pred_after = l2.forward(l1.forward(X_test))

r2_before = r2_score(y_test, y_pred_before)
r2_after = r2_score(y_test, y_pred_after)

assert r2_before < r2_after

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
y_pred = l2.forward(l1.forward(X_test))

r2_lumo_nn = r2_score(y_test[:, 0], y_pred[:, 0])
r2_homo_nn = r2_score(y_test[:, 1], y_pred[:, 1])
print(f'R2 LUMO: {r2_lumo_nn}\nR2 HOMO: {r2_homo_nn}')

##### DO NOT CHANGE #####
```

<!-- grade_id: beating_ridge_test, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: beating_ridge_test - possible points: 1

# Beating Ridge - 1 point

print(f'\tRidge\t\t\tNN\nLUMO\t{r2_lumo_ridge}\t{r2_lumo_nn}\nHOMO\t{r2_homo_ridge}\t{r2_homo_nn}')

assert r2_lumo_nn > r2_lumo_ridge and r2_homo_nn > r2_homo_ridge

##### DO NOT CHANGE #####
```

Quite possibly you were able to beat the ridge regression at that point, only with one hidden ReLu layer. Of course the hyperparameters of both models haven't been optimized yet, so you can get even better. But please do this in a separate notebook.

Let's have a look at the final plots:

```python
##### DO NOT CHANGE #####
fig, axs = plt.subplots(1, 2, figsize=(12, 6))
axs[0].scatter(y_test[:, 0], y_pred[:, 0], alpha=0.2, s=5)
axs[0].plot([-6, -1], [-6, -1], 'k')
axs[0].set_title(f'R2 LUMO: {r2_lumo_nn}')
axs[0].set_xlabel('true LUMO')
axs[0].set_ylabel('predicted LUMO')
axs[0].set_xlim([-6, -1])
axs[0].set_ylim([-6, -1])
axs[1].scatter(y_test[:, 1], y_pred[:, 1], alpha=0.2, s=5)
axs[1].plot([-9, -3], [-9, -3], 'k')
axs[1].set_title(f'R2 HOMO: {r2_homo_nn}')
axs[1].set_xlabel('true HOMO')
axs[1].set_ylabel('predicted HOMO')
axs[1].set_xlim([-9, -3])
axs[1].set_ylim([-9, -3])

plt.show()

##### DO NOT CHANGE #####
```